# Lab 2 


**Lab 2**
Assignment Objectives:  
Upon successful completion of this assignment, a student will be able to:
* Gain experience using and writing functions.
* Use basic statistical functions
* Gain experience in exploratory data visualization
* Use basic data cleanging.

The next code cell is required to use the autograder when running on Colab/DeepNote (uncomment the code). If you are using an environment with `otter-grader` already installed, then leave the code commented out.  

```python
#  Import libraries 
import pandas as pd
import numpy as np
import matplotlib as mpl
import matplotlib.pyplot as plt
%matplotlib inline
import seaborn as sns

import otter
grader = otter.Notebook()
```

*Tip: It is good practice to list all imports needed at the top of the notebook.  You can import modules in later cells as needed, but listing them at the top clearly shows all which are needed to be available / installed.*

### Data 

You will again be using the Auto-MPG data set available on the UCI repository. 

https://archive.ics.uci.edu/ml/datasets/Auto+MPG



## Exercise 1 - Import Data 
You will want to create a function to import the data 

The function should be named `ImportData(url, col_names)`    
`ImportData` has two input arguments: 
* `url` - a string that links to the data to be imported, and 
* `col_names` - a list of strings container each column names. 

In the function, you should import the data using `read_csv`, set the column names to `col_names`, and set the DataFrame index (*row labels*) to `carName`.  

The column names should be set to: 'mpg', 'cyl', 'disp', 'hp', 'weight', 'accel', 'model year', 'origin', 'car name'. 

For this function, do not use the `na_values` argument of `read_csv`.  We will handle missing data later in the notebook.

The function should return the DataFrame. 

Helpful links: 

* https://pandas.pydata.org/pandas-docs/stable/generated/pandas.read_csv.html#pandas.read_csv
* https://pandas.pydata.org/pandas-docs/stable/generated/pandas.read_table.html#pandas.read_table

```python
# Create the function ImportData here and return the data 

def ImportData(url, col_names):
    df = pd.read_csv(url,header = None, names = cNames,  delim_whitespace=True)
    df.set_index('car name', inplace = True,drop = False)
    
    # return the created DataFrame, df 
    
    
    return df

url = 'https://archive.ics.uci.edu/ml/machine-learning-databases/auto-mpg/auto-mpg.data'  
cNames = ['mpg', 'cyl', 'disp', 'hp', 'weight', 'accel', 'model year', 'origin', 'car name']
data = ImportData(url, cNames)
data.head()
```

Output:
```
mpg  cyl   disp     hp  weight  accel  model year  \
car name                                                                        
chevrolet chevelle malibu  18.0    8  307.0  130.0  3504.0   12.0          70   
buick skylark 320          15.0    8  350.0  165.0  3693.0   11.5          70   
plymouth satellite         18.0    8  318.0  150.0  3436.0   11.0          70   
amc rebel sst              16.0    8  304.0  150.0  3433.0   12.0          70   
ford torino                17.0    8  302.0  140.0  3449.0   10.5          70   

                           origin                   car name  
car name                                                      
chevrolet chevelle malibu       1  chevrolet chevelle malibu  
buick skylark 320               1          buick skylark 320  
plymouth satellite              1         plymouth satellite  
amc rebel sst                   1              amc rebel sst  
ford torino                     1                ford torino
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

## Exercise 2: Data Cleaning 

The data cleaning part will have multiple goals: to handle any missing data and to process the string data to remove spaces and replace with an underscore. 

You will create a function, named ColDataClean(data) that:

* Identify any Null/NA/NaN entries in the DataFrame, print out what entries will be removed, and remove them 
* Change the column labels (headers) from all lower case to [title case](https://danielmiessler.com/p/a-list-of-different-case-types/) and then replace "space" with underscore, "_". Print the column headers before and after.
* Change the DataFrame index (row labels) also be title case, that is, replace all spaces with underscore

```python
def ColDataClean(df):
    # Identify, print, and remove any entries with Null / NA / NaN
    print(df.isna().any())
  
    
    print("*** Removing null entries ***")
    df.dropna(axis = 0, how = 'any', inplace = True)
    
    
    print("*** String Operations on Column Headers ***")
    print("  Before: ", df.columns)
    df.columns = df.columns.str.title() 
    
    print("  After: ", df.columns)
    df.columns = df.columns.str.replace(' ', '_')
    
    print("*** String Operations on index strings *** \n")
    df.index = df.index.str.title()
    df.index = df.index.str.replace(' ', '_')
    
    
    return df

data = ColDataClean(data)

data.head()
```

Output:
```
mpg           False
cyl           False
disp          False
hp            False
weight        False
accel         False
model year    False
origin        False
car name      False
dtype: bool
*** Removing null entries ***
*** String Operations on Column Headers ***
  Before:  Index(['mpg', 'cyl', 'disp', 'hp', 'weight', 'accel', 'model year', 'origin',
       'car name'],
      dtype='object')
  After:  Index(['Mpg', 'Cyl', 'Disp', 'Hp', 'Weight', 'Accel', 'Model Year', 'Origin',
       'Car Name'],
      dtype='object')
*** String Operations on index strings *** 

                            Mpg  Cyl   Disp     Hp  Weight  Accel  Model_Year  \
car name                                                                        
Chevrolet_Chevelle_Malibu  18.0    8  307.0  130.0  3504.0   12.0          70   
Buick_Skylark_320          15.0    8  350.0  165.0  3693.0   11.5          70   
Plymouth_Satellite         18.0    8  318.0  150.0  3436.0   11.0          70   
Amc_Rebel_Sst          
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

## Exercise 3: Data Cleaning and Manipulation 

Look at the types for each column of the DataFrame. 

```python
data.dtypes
```

Output:
```
Mpg           float64
Cyl             int64
Disp          float64
Hp             object
Weight        float64
Accel         float64
Model_Year      int64
Origin          int64
Car_Name       object
dtype: object
```

Horsepower, `Hp`, is considered an object.  But, really this should be a float.  Why is this a float?

```python
data[data["Hp"] == '?']
```

Output:
```
Mpg  Cyl   Disp Hp  Weight  Accel  Model_Year  Origin  \
car name                                                                       
Ford_Pinto            25.0    4   98.0  ?  2046.0   19.0          71       1   
Ford_Maverick         21.0    6  200.0  ?  2875.0   17.0          74       1   
Renault_Lecar_Deluxe  40.9    4   85.0  ?  1835.0   17.3          80       2   
Ford_Mustang_Cobra    23.6    4  140.0  ?  2905.0   14.3          80       1   
Renault_18I           34.5    4  100.0  ?  2320.0   15.8          81       2   
Amc_Concord_Dl        23.0    4  151.0  ?  3035.0   20.5          82       1   

                                  Car_Name  
car name                                    
Ford_Pinto                      ford pinto  
Ford_Maverick                ford maverick  
Renault_Lecar_Deluxe  renault lecar deluxe  
Ford_Mustang_Cobra      ford mustang cobra  
Renault_18I                    renault 18i  
Amc_Concord_Dl              amc concord dl
```

There are some missing values in this column that need to be handled.

Create a function `HandleHP(df, col)`, that changes all the "?" values to a replacement value 99 and converts the column to be a float, where
* `df` is a DataFrame to be worked on 
* `col` is the name of the column to be "handled"



```python
def HandleHp(df, col):
    # Guarding 
    if (df[col].dtypes == 'object'):
       
        print ("  Current data type is : ", df[col].dtypes)
        df[col] = df[col].replace('?', 99).astype(float)
        print ("  Changed data type is : ", df[col].dtypes)
        
    return df 


HandleHp(data, 'Hp')
print(data.dtypes)
```

Output:
```
Current data type is :  object
  Changed data type is :  float64
Mpg           float64
Cyl             int64
Disp          float64
Hp            float64
Weight        float64
Accel         float64
Model_Year      int64
Origin          int64
Car_Name       object
dtype: object
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

## Exercise 4: Data Aggregation and Grouping 

Add two new columns for Kmpl (km per liter) and Kmpg (km per gallon). 

Note, the conversions are 1 mpg = 0.425144 kmpl  and 1 mpg = 1.609 kmpg.

_Type your answer here, replacing this text._

```python
#  Add two new columns
data['Kmpl'] =  data['Mpg']* 0.425144
data['Kmpg'] =  data['Mpg']* 1.609
print (data.head())

```

Output:
```
Mpg  Cyl   Disp     Hp  Weight  Accel  Model_Year  \
car name                                                                        
Chevrolet_Chevelle_Malibu  18.0    8  307.0  130.0  3504.0   12.0          70   
Buick_Skylark_320          15.0    8  350.0  165.0  3693.0   11.5          70   
Plymouth_Satellite         18.0    8  318.0  150.0  3436.0   11.0          70   
Amc_Rebel_Sst              16.0    8  304.0  150.0  3433.0   12.0          70   
Ford_Torino                17.0    8  302.0  140.0  3449.0   10.5          70   

                           Origin                   Car_Name      Kmpl    Kmpg  
car name                                                                        
Chevrolet_Chevelle_Malibu       1  chevrolet chevelle malibu  7.652592  28.962  
Buick_Skylark_320               1          buick skylark 320  6.377160  24.135  
Plymouth_Satellite              1         plymouth satellite  7.652592  28.962  
Amc_Rebel_Sst                   1              amc rebe
```

```python
# Add Region_Origin column 
region = {1:'ASIA', 2: 'EUROPE', 3: 'USA'}
data['Region_Origin'] = data['Origin'].map(region)


```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
```

## Exercise 5: Grouping 

Let's create a number of different groups for the data set. 

* Group 1 - group by 'region origin' 
* Group 2 - group by 'model year' 
* Group 3 - group by 'model year' and 'region origin' 

Let's look at these groups iterating with for loops

```python
# Create groups 

group1 = data.groupby(['Region_Origin'])
group2 = data.groupby(['Model_Year'])
group3 = data.groupby(['Model_Year','Region_Origin'])


# Print the data from Group 2 with a for loop 
#  restrict ourselves to only printing 4 cars in each group. 


    
for x, group in group2:
    print (x)
    print ("*** Start ")
    print (group.head(4))
    print ("*** End ")
    


```

Output:
```
(70,)
*** Start 
                            Mpg  Cyl   Disp     Hp  Weight  Accel  Model_Year  \
car name                                                                        
Chevrolet_Chevelle_Malibu  18.0    8  307.0  130.0  3504.0   12.0          70   
Buick_Skylark_320          15.0    8  350.0  165.0  3693.0   11.5          70   
Plymouth_Satellite         18.0    8  318.0  150.0  3436.0   11.0          70   
Amc_Rebel_Sst              16.0    8  304.0  150.0  3433.0   12.0          70   

                           Origin                   Car_Name      Kmpl  \
car name                                                                 
Chevrolet_Chevelle_Malibu       1  chevrolet chevelle malibu  7.652592   
Buick_Skylark_320               1          buick skylark 320  6.377160   
Plymouth_Satellite              1         plymouth satellite  7.652592   
Amc_Rebel_Sst                   1              amc rebel sst  6.802304   

                             Kmpg Region_Origin  
c
```

```python
grader.check("q5")
```

Output:
```
q5 results: All test cases passed!
```

<!-- BEGIN QUESTION -->

## Exercise 6: Data Exploration and Analysis 

Create a function `plot_mpg_region()` that reports the mpg over the grouping `Region_Origin`.  Consider using line, bar, scatter plots.  

The function should: 

* take as input a DataFrame group (from above)
* report mpg over the group 
* make sure the plot looks good and has labels, title, etc. 

```python
def plot_mpg_region(my_group):

    avg_mpg_by_region = my_group.groupby(['Region_Origin']).mean()
    #Jupyter notebooks have a default figure size, which might override the figsize parameter. 
    #To change the figure size in a Jupyter notebook, we need to use %matplotlib inline and rcParams.
    avg_mpg_by_region.plot.scatter(x="pH", y="alcohol")
    

    return 


plot_mpg_region(group1)
```

Output:
```
Cell In[62], line 3, in plot_mpg_region(my_group)
      1 def plot_mpg_region(my_group):
----> 3     avg_mpg_by_region = my_group.pd.groupby(['Region_Origin']).mean()
      4     #Jupyter notebooks have a default figure size, which might override the figsize parameter. 
      5     #To change the figure size in a Jupyter notebook, we need to use %matplotlib inline and rcParams.
      6     avg_mpg_by_region.plot.scatter(x="pH", y="alcohol")

File ~/.conda/envs/un5550/lib/python3.10/site-packages/pandas/core/groupby/groupby.py:952, in GroupBy.__getattr__(self, attr)
    949 if attr in self.obj:
    950     return self[attr]
--> 952 raise AttributeError(
    953     f"'{type(self).__name__}' object has no attribute '{attr}'"
    954 )

AttributeError: 'DataFrameGroupBy' object has no attribute 'pd'
```

```python
def plot_mpg_region(my_group):

    avg_mpg_by_region = data.groupby(['Mpg']).mean()
    #Jupyter notebooks have a default figure size, which might override the figsize parameter. 
    #To change the figure size in a Jupyter notebook, we need to use %matplotlib inline and rcParams.
    plt.rcParams["figure.figsize"] = (16, 6)
    avg_mpg_by_region.plot(kind = 'bar', color = ['yellow','black','pink'])
    plt.title(' Report of Avg Mpg over the Group Region_Origin')
    plt.xlabel('Region_Origin')
    plt.ylabel('Mpg')
    plt.xticks(rotation = 45)
    plt.grid(axis = 'y', linestyle = '--',alpha = 0.7)
    

    return 


plot_mpg_region(group1)
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

## Exercise 7: Data Exploration and Analysis 

Create a function `plot_mpg_year()` that reports the mpg over the grouping `Model Year`.  Consider using line, bar, scatter plots.  

The function should: 

* take as input a DataFrame group (from above)
* report mpg over the group 
* make sure the plot looks good and has labels, title, etc. 

```python
def plot_mpg_year(my_group):
    
    
    return 


plot_mpg_year(group2)
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

## Exercise 8: Data Exploration and Analysis 

Create a function `plot_mpg_year_region()` that reports the mpg over the both groupings `Model Year` and `Region`.  Think about what is the appropriate type of plot.  

The function should: 

* take as input the DataFrame (*note, this is different than the other plots*)
* report mpg over both groups
* make sure the plot looks good and has labels, title, etc. 

```python
def plot_mpg_year_region(data):
    
    
    Mpg_pivot = data.pivot_table(values = 'Mpg', index = 'Model_Year', columns = 'Region_Origin' , aggfunc = 'mean')
    plt.rcParams["figure.figsize"] = (10, 6)
    sns.heatmap(Mpg_pivot, annot = True, cmap = 'coolwarm', fmt = '.2f', cbar = True, linewidth = 0.5)
    plt.title(' Mpg by the Model Year and Region_Origin')
    plt.xlabel('Region_Origin')
    plt.ylabel('Model_Year')
    plt.tight_layout()
    plt.show()
    return 


plot_mpg_year_region(data)
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

## Exercise 9: Explore Additional aspects of the data using statistics 

Look at the statistics for 2 of the `float` variables in the data set.  
Report the mean, standard deviation, 5-number summary.  

Plot the distributions using histograms.  

Which distribution has the most skew? 

Look at the statistics and distribution again now grouped by country of origin.  

Look at the statistics and distribution again now grouped by number of cylinders.  

**Look at the statistics for 2 of the float variables in the data set.
Report the mean, standard deviation, 5-number summary.** 

```python

#  I am selecting Mpg and Acceleration as my two float varaibles because both of them as more float values 
# For Mpg Stat
Mpg_Mean =  np.mean(data['Mpg'])
Mpg_Standard_Deviation =  np.std(data['Mpg'])
Mpg_5_Number_Summary = [np.max(data['Mpg']),np.min(data['Mpg']),np.median(data['Mpg']),np.percentile(data['Mpg'],25),np.percentile(data['Mpg'],75)]
Max = Mpg_5_Number_Summary[0]
Min = Mpg_5_Number_Summary[1]
Median = Mpg_5_Number_Summary[2]
Percentile_at_25 = Mpg_5_Number_Summary[3]
Percentile_at_75 = Mpg_5_Number_Summary[4]

#printing Mpg Statistics
print("Statistics for 'Mpg': ")
print('Mean :',Mpg_Mean)
print('Standard Deviation:', Mpg_Standard_Deviation)
print('5-number summary ',Mpg_5_Number_Summary)
print(f'Maximum : {Max}, Minimum : {Min}, Medium : {Median},Percentile_at_25 : {Percentile_at_25}, Percentile_at_75 : {Percentile_at_75}')

# For Acceleration Statistics
Acceleration_Mean =  np.mean(data['Accel'])
Acceleration_Standard_Deviation =  np.std(data['Accel'])
Acceleration_5_Number_Summary = [np.max(data['Accel']),np.min(data['Accel']),np.median(data['Accel']),np.percentile(data['Mpg'],25),np.percentile(data['Accel'],70)]
Max = Acceleration_5_Number_Summary[0]
Min = Acceleration_5_Number_Summary[1]
Median = Acceleration_5_Number_Summary[2]
Percentile_at_25 = Acceleration_5_Number_Summary[3]
Percentile_at_75 = Acceleration_5_Number_Summary[4]

#printing Acceleration Stat
print("\nStatistics for 'Acceleration': ")
print('Mean :',Acceleration_Mean)
print('Standard Deviation:', Acceleration_Standard_Deviation)
print('5-number summary :',Acceleration_5_Number_Summary)
print(f'Maximum : {Max}, Minimum : {Min}, Medium : {Median},Percentile_at_25 : {Percentile_at_25}, Percentile_at_75 : {Percentile_at_75}')

```

**Plot the distributions using histograms.**

```python
#Histogram for Mpg
plt.rcParams["figure.figsize"] = (10, 6)
plt.hist(data['Mpg'],bins = 30, color = 'yellow', edgecolor = 'black')
plt.title('Distribution of Mpg')
plt.xlabel('Mpg')
plt.ylabel('Frequency')
plt.tight_layout()
plt.show()
```

```python
#Histogram for Acceleration
plt.rcParams["figure.figsize"] = (10, 6)
plt.hist(data['Accel'],bins = 30, color = 'skyblue', edgecolor = 'black')
plt.title('Distribution of Acceleration')
plt.xlabel('Mpg')
plt.ylabel('Frequency')
plt.tight_layout()
plt.show()
```

Q)Which distribution has the most skew? 
<br>Ans: As we can the distribution of both Mpg and Acceleration. I would say the distribution of **Mpg is most skewed**
although it is postively skewed. 
<br> We can obeserve the values of Mean, Median and Mode of the both Mpg and Acceleration. When you plot those values for both Mpg and Acceleration we can see that Acceleration is kinda of symmetrical.

**Look at the statistics and distribution again now grouped by country of origin.** 

```python
# Calculating Statistics and Distribution agian by grouping country of origin for Mpg
Grouped_By_Region_Origin = data.groupby('Region_Origin')
Mpg_Statistics_By_Country_Of_Origin = Grouped_By_Region_Origin['Mpg'].describe()
Mpg_Statistics_By_Country_Of_Origin
```

```python
# Calculating Statistics and Distribution agian by grouping country of origin for Acceleration
Acceleration_Statistics_By_Country_Of_Origin = Grouped_By_Region_Origin['Accel'].describe()
Acceleration_Statistics_By_Country_Of_Origin
```

**Look at the statistics and distribution again now grouped by number of cylinders.**


```python
Grouped_By_Cylinders = data.groupby('Cyl')
Mpg_Statistics_By_Grouping_Cylinders = Grouped_By_Cylinders['Mpg'].describe()
Mpg_Statistics_By_Grouping_Cylinders
```

```python
Grouped_By_Cylinders = data.groupby('Cyl')
Acceleration_Statistics_By_Grouping_Cylinders = Grouped_By_Cylinders['Accel'].describe()
Acceleration_Statistics_By_Grouping_Cylinders
```

<!-- END QUESTION -->



## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: [2, 2, 2])
```
>>> data.iloc[0,0] == 18.0 and data.iloc[0,8] == 'chevrolet chevelle malibu'
True
```
```
>>> data.iloc[4,0] == 17.0 and data.iloc[4,8] == 'ford torino'
True
```
```
>>> data.shape[0]==398 and data.shape[1]==9
True
```


Tests for q2 (points: [2, 2, 2])
```
>>> data.isnull().values.any() == False
True
```
```
>>> data.columns[6] =='Model_Year'
True
```
```
>>> data.index[0] =='Chevrolet_Chevelle_Malibu' and data.index[4] =='Ford_Torino'
True
```


Tests for q3 (points: [2, 2])
```
>>> data["Hp"].values.any() != '?'
True
```
```
>>> data['Hp'].dtypes=='float64'
True
```


Tests for q4 (points: [2, 2])
```
>>> data.iloc[0,9] == 7.652592 and data.iloc[0,10] == 28.962
True
```
```
>>> data.iloc[1,9] == 6.377160 and np.isclose(data.iloc[1,10], 24.135)
True
```


Tests for q5 (points: [4])
```
>>> total_gp = [x[0] for x, group in group2]
>>> total_gp == [70, 71, 72, 73, 74, 75, 76, 77, 78, 79, 80, 81, 82]
True
```
