# Lab 5 - Getting Data

The topics of week 5 is getting data with web scraping. 

In this lab notebook you will gain experience webscraping data of various forms.

We will be looking at various websites and using them to get information (these sites are selected because of their public informative nature). 

## Lab Setup 

```python
from bs4 import BeautifulSoup
import requests 
import pandas as pd
import numpy as np
import re

import otter
grader = otter.Notebook()
```

Output:
```
File ~/.conda/envs/un5550/lib/python3.10/site-packages/otter/check/notebook.py:153, in Notebook._log_event(self, event_type, results, question, success, error, shelve_env)
    144 if _SHELVE and event_type == EventType.CHECK:
    145     entry.shelve(
    146         shelve_env,
    147         delete=True,
   (...)
    150         variables=self._vars_to_store
    151     )
--> 153 entry.flush_to_file(_OTTER_LOG_FILENAME)

File ~/.conda/envs/un5550/lib/python3.10/site-packages/otter/check/logs.py:139, in LogEntry.flush_to_file(self, filename)
    136     dill.dump(self, file)
    138 except OSError:
--> 139     raise Exception(
    140         "Could not create the log file as the file system is read-only. Please contact your "
    141         "instructor before continuing on this assignment."
    142     )
    144 finally:
    145     if "file" in locals():

Exception: Could not create the log file as the file system is read-only. Please contact your instructor before continuing on t
```

## Exercise 1 

Identify the name of the data set most relevant to 'climate' on https://www.data.gov 

You will want to use the **url = "https://catalog.data.gov/dataset?q=climate&sort=score+desc%2C+name+asc"**.

Then, you will use BeautifulSoup to parse the website to find the name of the data set.  

```python
# Identify the name of the data set most relevant to "climate" on data.gov
# Use BeautifulSoup to parse the website to find the name of the data set. 
url = 'https://catalog.data.gov/dataset?q=climate&sort=score+desc%2C+name+asc'
resp_q1 = requests.get(url) 
soup_q1 = BeautifulSoup(resp_q1,'html.parser')

#using header tag to get the text 
name = soup_q1.h3.a.text.strip()
name
```

Output:
```
Cell In[18], line 5
      3 url = 'https://catalog.data.gov/dataset?q=climate&sort=score+desc%2C+name+asc'
      4 resp_q1 = requests.get(url) 
----> 5 soup_q1 = BeautifulSoup(resp_q1,'html.parser')
      7 #using header tag to get the text 
      8 name = soup_q1.h3.a.text.strip()

File ~/.conda/envs/un5550/lib/python3.10/site-packages/bs4/__init__.py:315, in BeautifulSoup.__init__(self, markup, features, builder, parse_only, from_encoding, exclude_encodings, element_classes, **kwargs)
    313 if hasattr(markup, 'read'):        # It's a file-type object.
    314     markup = markup.read()
--> 315 elif len(markup) <= 256 and (
    316         (isinstance(markup, bytes) and not b'<' in markup)
    317         or (isinstance(markup, str) and not '<' in markup)
    318 ):
    319     # Issue warnings for a couple beginner problems
    320     # involving passing non-markup to Beautiful Soup.
    321     # Beautiful Soup will still parse the input as markup,
    322     # since that is som
```

```python
grader.check("q1")
```

Output:
```
File ~/.conda/envs/un5550/lib/python3.10/site-packages/otter/check/notebook.py:153, in Notebook._log_event(self, event_type, results, question, success, error, shelve_env)
    144 if _SHELVE and event_type == EventType.CHECK:
    145     entry.shelve(
    146         shelve_env,
    147         delete=True,
   (...)
    150         variables=self._vars_to_store
    151     )
--> 153 entry.flush_to_file(_OTTER_LOG_FILENAME)

File ~/.conda/envs/un5550/lib/python3.10/site-packages/otter/check/logs.py:139, in LogEntry.flush_to_file(self, filename)
    136     dill.dump(self, file)
    138 except OSError:
--> 139     raise Exception(
    140         "Could not create the log file as the file system is read-only. Please contact your "
    141         "instructor before continuing on this assignment."
    142     )
    144 finally:
    145     if "file" in locals():

Exception: Could not create the log file as the file system is read-only. Please contact your instructor before continuing on t
```

## Exercise 2 

Total number of clinical trials ("research studies") recorded at the National Institutes of Health, https://clinicaltrials.gov

```python
# Identify the number of clinical trials (research studies) on https://clinicaltrials.gov
# Use BeautifulSoup to parse the website to find the number of studies. 
url = 'https://classic.clinicaltrials.gov/'
resp_q2 = requests.get(url) 
soup_q2 = BeautifulSoup(resp_q2,'html.parser')

#using header tag to get the text and converting to integer.
num_research_studies =  int(soup_q2.h2.text.split()[1].replace(',','')) # integer value
num_research_studies
```

```python
grader.check("q2")
```

## Exercise 3 

Give the number of travel advisories, by level, for international travel given by the US government. Data available at: https://travel.state.gov/content/travel/en/traveladvisories/traveladvisories.html/

Create two solutions for this question, the first using `pandas` `read_html` function.  The second version should be using just BeautifulSoup and Python to scrape the table. 

Note, you do not need to put the results into a DataFrame for version 2 (although that is one option). 

```python
site = requests.get('https://travel.state.gov/content/travel/en/traveladvisories/traveladvisories.html/').text
q3 = BeautifulSoup(site, 'html5lib')
```

**Version 1** - Write one solution using pandas `read_html`

**Version 2** - Write another solution just using beautiful soup and other Python functions to scrape the table.

```python
# Version 1
# Read in the table using pandas "read_html" function 
# Report out number of each warnings 


travel_advisory = pd.read_html("https://travel.state.gov/content/travel/en/traveladvisories/traveladvisories.html/")[0]
travel_advisory_level_counts = travel_advisory['Level'].value_counts()

l1Num1 = travel_advisory_level_counts.get('Level 1: Exercise Normal Precautions')
l2Num1 = travel_advisory_level_counts.get('Level 2: Exercise Increased Caution')
l3Num1 = travel_advisory_level_counts.get('Level 3: Reconsider Travel')
l4Num1 = travel_advisory_level_counts.get('Level 4: Do Not Travel')




print("Number of travel Level 1 warnings : " +  str(l1Num1)) 
print("Number of travel Level 2 warnings : " +  str(l2Num1))
print("Number of travel Level 3 warnings : " +  str(l3Num1))
print("Number of travel Level 4 warnings : " +  str(l4Num1))
```

```python
# Version 2
# Write another solution just using beautiful soup and other Python functions to 
#   scrape the table information
# Report out number of each warnings 


# for getting the table from website you can find the table tag and has id summary using it we can get the table.
travel_advisories_table = q3.find('table', {'summary': ''})


#creating two lists for storing the values
advisories = []
levels = []

#for storing all the data from the table. 
for row in travel_advisories_table.find_all('tr')[1:]:
    columns = row.find_all('td')
    advisories.append(columns[0].text.strip())
    levels.append(columns[1].text.strip())

# Creating a dataframe with same coulmn names as in the website 
travel_advisory = pd.DataFrame({'Advisory': advisories, 'Level': levels})


#using same code above to get the level values
travel_advisory_level_counts = travel_advisory['Level'].value_counts()

l1Num2 = travel_advisory_level_counts.get('Level 1: Exercise Normal Precautions')
l2Num2 = travel_advisory_level_counts.get('Level 2: Exercise Increased Caution')
l3Num2 = travel_advisory_level_counts.get('Level 3: Reconsider Travel')
l4Num2 = travel_advisory_level_counts.get('Level 4: Do Not Travel')



print("Number of travel Level 1 warnings : " +  str(l1Num2)) 
print("Number of travel Level 2 warnings : " +  str(l2Num2))
print("Number of travel Level 3 warnings : " +  str(l3Num2))
print("Number of travel Level 4 warnings : " +  str(l4Num2))
```

```python
grader.check("q3")
```

##  Exercise 4

Let's look at the number of workplace fatalities as reported to the federal and state Occupational Safety and Health Administration (OSHA) at https://www.osha.gov/fatalities. 

Use `pandas` `read_html` to read in the table of fatalities. 


```python
headers = {"User-Agent":"Mozilla/5.0"}
site = requests.get('https://www.osha.gov/fatalities', headers = headers).text

soup_q4 = BeautifulSoup(site,'html.parser') 

# Read in Table using pandas "read_html" method
df4 = pd.read_html(site)
df4 = df4[0]

# Report the total number of citations issued related to fatality
cits = (df4['Citation Issued Related to Fatality'] == 'Yes').sum()

df4.head()

print("Number of citations issues related to fatality: ", cits)
```

```python
grader.check("q4")
```

## Exercise 5 

Save in a DataFrame the incidents that occurred in fiscal year 2022, Oct. 1, 2021 - Sept. 30, 2022.  

Try using `pd.to_datetime` to add a new "Date" column with the Date of Incident column converted to a datetime object that may be easier to filter. 

https://pandas.pydata.org/docs/reference/api/pandas.to_datetime.html


```python
# Save all incidents that occured in fiscal year 2022: Oct. 1, 2021 - Sept. 30, 2022
#  in a DataFrame ("q5result"), it will have a new column "Date" with the Date of 
#  Incident column converted to a datetime object. 


df4['Date'] = pd.to_datetime(df4['Date of Incident'], format='%m/%d/%Y')
starting_date = pd.to_datetime('2021-10-01')
ending_date = pd.to_datetime('2022-09-30')


q5result = df4[(df4['Date'] >= starting_date) & (df4['Date'] <= ending_date)]
q5result.head()
```

```python
grader.check("q5")
```

## Exercise 6 

Identify the most viewed data set on Michigan state's open data portal.   
https://data.michigan.gov/browse?sortBy=most_accessed

Stores its name in the provided "most_viewed_data_nm" variable and the number of views in "num_views".

```python
resp_q6 = requests.get('https://data.michigan.gov/browse?sortBy=most_accessed').text 
soup_q6 = BeautifulSoup(resp_q6,'html.parser')

most_viewed_data_nm = soup_q6.find('div', {'class': 'browse2-result-title'}).a.text

num_views = int(float(soup_q6.find('div',{'class':'browse2-result-view-count-value'}).text.strip().replace(',', '')))


print("Most visited data: %s  with %d views" % (most_viewed_data_nm, num_views))
```

```python
grader.check("q6")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1 (points: 6)
```
>>> name == "Downscaled Climate Model Climate Toolbox"
True
```


Tests for q2 (points: 8)
```
>>> num_research_studies >= 468124 
True
```


Tests for q3 (points: 10)
```
>>> l1Num1 == l1Num2 
True
```
```
>>> l2Num1 == l2Num2 
True
```
```
>>> l3Num1 == l3Num2
True
```
```
>>> l4Num1 == l4Num2
True
```


Tests for q4 (points: 10)
```
>>> df4.shape[1] == 7
True
```


Tests for q5 (points: 6)
```
>>> q5result.shape[0] == 864
True
```
```
>>> q5result.shape[1] == 8
True
```
```
>>> "Date" in q5result.columns
True
```


Tests for q6 (points: 10)
```
>>> most_viewed_data_nm == 'MI_Congress_Plot'
True
```
